# Week 1: Python Fundamentals for AI & Machine Learning

Welcome to your first Python programming session!

## What is Python?

Python is a **programming language** – a way for humans to give instructions to computers. Think of it like learning a new language to communicate, but instead of talking to people, you're talking to machines!

<div dir="rtl">

<div dir="rtl">

### ما هي لغة Python؟

</div>

<div dir="rtl">

لغة <code>Python</code> هي <strong>لغة برمجة</strong> – أي طريقة يستخدمها البشر لإعطاء الأوامر للكمبيوتر. تخيل الموضوع مثل تعلم لغة جديدة للتواصل، لكن بدلاً من التحدث مع الناس، أنت تتواصل مع الآلات!

</div>

</div>

---

## How Does Code Become Actions?

When you write Python code, here's the journey it takes:

```
┌─────────────────┐
│  You write      │
│  Python code    │  ← Human-readable
└────────┬────────┘
         │
         ▼
┌─────────────────┐
│  Python         │
│  Interpreter    │  ← Translates your code to bytecode
└────────┬────────┘
         │
         ▼
┌─────────────────┐
│  Bytecode       │  ← Intermediate form (.pyc files)
└────────┬────────┘
         │
         ▼
┌─────────────────┐
│  Python Virtual │
│  Machine (PVM)  │  ← Executes bytecode instructions
└────────┬────────┘
         │
         ▼
┌─────────────────┐
│  Computer       │
│  executes!      │  ← Actions happen!
└─────────────────┘
```

<div dir="rtl">

<div dir="rtl">

### كيف يصبح الـ <code>code</code> أوامر تُنفذ؟

</div>

<div dir="rtl">

عندما تكتب <code>code</code> بلغة <code>Python</code>، الكمبيوتر يترجمها على مراحل:

</div>

<ul>
<div dir="rtl">

  <li>أول شيء الـ <code>interpreter</code> يحول الكود الذي كتبته (مفهوم للبشر) إلى <code>bytecode</code> (لغة وسيطة).</li>
  <li>ثم يقوم الـ <code>Python Virtual Machine</code> بتنفيذ هذه التعليمات.</li>
  <li>وفي النهاية، ينفذ الكمبيوتر الأوامر فعليًا!</li>

</div>
</ul>

</div>

---

## What You'll Learn This Week

This is a **PRACTICAL** week – you'll learn by experimenting and observing!

1. **Variables & Data Types** – Understanding different kinds of data
2. **Lists & Indexing** – Organizing multiple values
3. **Loops & Iteration** – Repeating actions automatically
4. **Functions** – Creating reusable code blocks

<div dir="rtl">

<div dir="rtl">

### ماذا ستتعلم هذا الأسبوع

</div>

<div dir="rtl">

هذا أسبوع <strong>عملي</strong> – ستتعلم من خلال التجربة والملاحظة!

</div>

<ol>
<div dir="rtl">

  <li><code>Variables</code> و <code>Data Types</code> – فهم الأنواع المختلفة من البيانات</li>
  <li><code>Lists</code> و <code>Indexing</code> – تنظيم القيم المتعددة</li>
  <li><code>Loops</code> و <code>Iteration</code> – تكرار العمليات تلقائيًا</li>
  <li><code>Functions</code> – إنشاء قطع <code>code</code> قابلة لإعادة الاستخدام</li>

</div>
</ol>

</div>

---

**Estimated Time:** 3–4 hours  
Let’s begin! 🚀


---

## Part 1: Variables and Data Types

### What is a Variable?

A **variable** is like a labeled box that stores a value. You give it a name, and Python remembers what's inside.

```
┌─────────────────┐
│  Box labeled    │
│  "age"          │  ← The name we give
│                 │
│  Contains: 25   │  ← The value stored
└─────────────────┘
```

In Python: `age = 25`

<div dir="rtl">

### ما هو الـ Variable؟

</div>

<div dir="rtl">

الـ **variable** مثل صندوق عليه اسم بيخزن قيمة معينة. أنت تدي له
اسم، و Python بيتذكر
ما بداخله.

</div>

<div dir="rtl">

في Python نكتب: `age = 25`

</div>

### Let's Create Our First Variable!


```python
# Run this cell! (Click on it and press Shift+Enter)
# The "#" symbol means this is a comment - Python ignores it

my_first_variable = 42

age = 25
print(age)

print(my_first_variable)
```

Output:
```
25
42
```

**What just happened?**

1. We created a variable called `my_first_variable`
2. We put the number `42` inside it
3. `print()` shows us what's inside

**Important**: The `=` sign in programming means "put this value on the right side into this variable on the left side" (not "equals" like in math!)

<div dir="rtl">

### ماذا حدث الآن؟

</div>

<div dir="rtl">

1. أنشأنا variable اسمه `my_first_variable`
2. وضعنا الرقم `42` بداخله
3. الـ `print()` يعرض لنا ما بداخل الـ variable

</div>

<div dir="rtl">

**مهم**: علامة `=` في البرمجة معناها "ضع القيمة اللي على اليمين داخل الـ variable اللي على اليسار" (مش معناها يساوي مثل الرياضيات!)

</div>

---

## Data Types: Different Kinds of Information

Just like in real life, data comes in different types. Let's explore each one by **experimenting**!

<div dir="rtl">

### أنواع الـ Data: أنواع مختلفة من المعلومات

</div>

<div dir="rtl">

تماماً مثل الحياة الواقعية، البيانات تأتي بأنواع مختلفة. لنستكشف كل نوع من خلال **التجربة**!

</div>


---

### 1. Integer (`int`) - Whole Numbers

**Integers** are whole numbers with **no decimal point**.

Examples: `1`, `2`, `100`, `-5`, `0`

Think of counting on your fingers: 1, 2, 3, 4, 5... (no fractions!)

<div dir="rtl">

### 1. Integer (`int`) - أرقام صحيحة

</div>

<div dir="rtl">

الـ **Integer** هي أرقام كاملة **بدون فاصلة عشرية**.

</div>

<div dir="rtl">

أمثلة: `1`, `2`, `100`, `-5`, `0`

</div>

<div dir="rtl">

تخيل العد على أصابعك: 1، 2، 3، 4، 5... (بدون كسور!)

</div>

Let's experiment with integers:


```python
# Different integers
positive_number = 10
negative_number = -5
zero = 0

print("Positive:", positive_number)
print("Negative:", negative_number)
print("Zero:", zero)
```

Output:
```
Positive: 10
Negative: -5
Zero: 0
```

```python
# Let's check the type
my_number = 25
print("Value:", my_number)
print("Type:", type(my_number).__name__)

# type() is a function that tells us what kind of data we have
```

Output:
```
Value: 25
Type: int
```

```python
# Math with integers
a = 10
b = 3

print("Addition:", a + b)
print("Subtraction:", a - b)
print("Multiplication:", a * b)
print("Division:", a / b)  # Notice: division gives a decimal!
```

Output:
```
Addition: 13
Subtraction: 7
Multiplication: 30
Division: 3.3333333333333335
```

**Key Observation**: When we divide integers, Python gives us a decimal number (more on this next!)

**Try it yourself**: Change the values of `a` and `b` in the cell above and run it again!

<div dir="rtl">

### ملاحظة مهمة

</div>

<div dir="rtl">

عندما نقسم integers، Python يعطينا رقم عشري (سنتحدث عن هذا بعد قليل!)

</div>

<div dir="rtl">

**جربها بنفسك**: غيّر قيم `a` و `b` في الـ cell أعلاه وشغلها مرة أخرى!

</div>


---

### 2. Float (`float`) - Decimal Numbers

**Floats** are numbers with **decimal points**.

Examples: `3.14`, `0.5`, `2.0`, `-1.75`, `0.001`

Think of measuring things precisely: height (1.75 meters), weight (68.5 kg), temperature (36.6°C)

**The key difference from integers**: THE DECIMAL POINT!

<div dir="rtl">

### 2. Float (`float`) - أرقام عشرية

</div>

<div dir="rtl">

الـ **Float** هي أرقام **بفاصلة عشرية**.

</div>

<div dir="rtl">

أمثلة: `3.14`, `0.5`, `2.0`, `-1.75`, `0.001`

</div>

<div dir="rtl">

تخيل قياس الأشياء بدقة: الطول (1.75 متر)، الوزن (68.5 كيلو)، درجة الحرارة (36.6 درجة)

</div>

<div dir="rtl">

**الفرق الأساسي عن الـ integer**: الفاصلة العشرية!

</div>

Let's see floats in action:


```python
# Different floats
pi = 3.14159
half = 0.5
tiny_number = 0.001
two_as_float = 2.0  # Has a decimal point, so it's a float!

print("Pi:", pi)
print("Half:", half)
print("Tiny:", tiny_number)
print("Two (float):", two_as_float)
```

Output:
```
Pi: 3.14159
Half: 0.5
Tiny: 0.001
Two (float): 2.0
```

```python
# Compare: Integer 2 vs Float 2.0
integer_two = 2
float_two = 2.0

print("Integer 2:")
print("  Value:", integer_two)
print("  Type:", type(integer_two))

print("\nFloat 2.0:")
print("  Value:", float_two)
print("  Type:", type(float_two))

print('\nThey look the same, but Python treats them differently!')
```

Output:
```
Integer 2:
  Value: 2
  Type: <class 'int'>

Float 2.0:
  Value: 2.0
  Type: <class 'float'>

They look the same, but Python treats them differently!
```

```python
# Math with floats
height = 1.75  # meters
weight = 68.5  # kg

print("Height:", height, "m")
print("Weight:", weight, "kg")
print("Height + Weight:", height + weight)
print("Height × 2:", height * 2)
```

Output:
```
Height: 1.75 m
Weight: 68.5 kg
Height + Weight: 70.25
Height × 2: 3.5
```

```python
# Floats are VERY important in machine learning!
# Here are some common ML values:

learning_rate = 0.01    # How fast a model learns (small decimal)
accuracy = 0.95         # Model accuracy: 95% = 0.95
loss = 0.234            # Model error

print("Learning rate:", learning_rate)
print("Accuracy:", accuracy)
print("Loss:", loss)

print("\nNotice: All these ML values are floats!")
```

Output:
```
Learning rate: 0.01
Accuracy: 0.95
Loss: 0.234

Notice: All these ML values are floats!
```

**Key Takeaway**:

- **Integers** = whole numbers: `1, 2, 3, 100`
- **Floats** = decimal numbers: `1.5, 2.0, 0.001, 3.14159`

In machine learning, we use floats A LOT because we need precise measurements!

<div dir="rtl">

### الخلاصة المهمة

</div>

<div dir="rtl">

- **Integer** = أرقام صحيحة: `1, 2, 3, 100`
- **Float** = أرقام عشرية: `1.5, 2.0, 0.001, 3.14159`

</div>

<div dir="rtl">

في machine learning، نستخدم floats كثيراً جداً لأننا نحتاج قياسات دقيقة!

</div>


---

### 3. String (`str`) - Text

**Strings** are text - letters, words, sentences.

**Important**: Strings must be wrapped in quotes: `"like this"` or `'like this'`

Examples: `"hello"`, `"Python"`, `"My name is Mahmoud"`, `"123"` (yes, numbers in quotes are strings!)

<div dir="rtl">

### 3. String (`str`) - نص

</div>

<div dir="rtl">

الـ **String** هي نص - حروف، كلمات، جمل.

</div>

<div dir="rtl">

**مهم**: الـ string يجب أن تكون داخل علامات تنصيص: `"هكذا"` أو `'هكذا'`

</div>

<div dir="rtl">

أمثلة: `"hello"`, `"Python"`, `"My name is Mahmoud"`, `"123"` (نعم، الأرقام داخل علامات تنصيص هي string!)

</div>

Let's play with strings:


```python
# Different strings
name = "Yousef"
greeting = 'Hello!'  # Single or double quotes both work
sentence = "Python is awesome"
number_as_text = "123"  # This is text, not a number!

print(name)
print(greeting)
print(sentence)
print(number_as_text)
```

Output:
```
Yousef
Hello!
Python is awesome
123
```

```python
# Check types - see the difference!
actual_number = 123
number_as_string = "123"

print("Number 123:")
print("  Value:", actual_number)
print("  Type:", type(actual_number))

print("\nString '123':")
print("  Value:", number_as_string)
print("  Type:", type(number_as_string))

print("\nThey look the same but are COMPLETELY different!")
```

Output:
```
Number 123:
  Value: 123
  Type: <class 'int'>

String '123':
  Value: 123
  Type: <class 'str'>

They look the same but are COMPLETELY different!
```

```python
# You can add strings together (called "concatenation")
first_name = "Ahmed"
last_name = "Ali"

full_name = first_name + " " + last_name
print(full_name)
```

Output:
```
Ahmed Ali
```

```python
# Strings in machine learning
model_name = "Neural Network"
dataset = "MNIST"
file_path = "data/images/cat.jpg"
class_label = "dog"

print("Model:", model_name)
print("Dataset:", dataset)
print("File path:", file_path)
print("Class:", class_label)

print("\nAll these are strings (text data)!")
```

Output:
```
Model: Neural Network
Dataset: MNIST
File path: data/images/cat.jpg
Class: dog

All these are strings (text data)!
```

**Key Takeaway**:

- Strings = text (always in quotes)
- `"123"` is different from `123` (string vs integer)
- Used for names, labels, file paths in ML

<div dir="rtl">

### الخلاصة المهمة

</div>

<div dir="rtl">

- String = نص (دائماً داخل علامات تنصيص)
- `"123"` يختلف عن `123` (string مقابل integer)
- نستخدمها للأسماء، الـ labels، مسارات الـ files في الـ machine learning

</div>


---

### 4. Boolean (`bool`) - True or False

**Booleans** have only TWO possible values: `True` or `False`

Think of it like:

- A light switch: ON or OFF
- A yes/no question
- A checkbox: checked or unchecked

**Important**: `True` and `False` are special words in Python (notice the capital letters!)

<div dir="rtl">

### 4. Boolean (`bool`) - صح أو خطأ

</div>

<div dir="rtl">

الـ **Boolean** لها قيمتان فقط: `True` أو `False`

</div>

<div dir="rtl">

تخيلها مثل:

</div>

<div dir="rtl">

- مفتاح الإضاءة: مشغّل أو مطفأ
- سؤال نعم/لا
- checkbox: محدد أو غير محدد

</div>

<div dir="rtl">

**مهم**: `True` و `False` كلمات خاصة في Python (لاحظ الحرف الكبير في البداية!)

</div>

Let's see booleans:


```python
# Simple booleans
is_sunny = True
is_raining = False

print("Is it sunny?", is_sunny)
print("Is it raining?", is_raining)
print("Type:", type(is_sunny))
```

Output:
```
Is it sunny? True
Is it raining? False
Type: <class 'bool'>
```

```python
# Booleans come from comparisons
age = 25

is_adult = age >= 18
is_child = age < 13
is_exactly_25 = age == 25  # Note: == for comparison, = for assignment

print("Age:", age)
print("Is adult (≥18)?", is_adult)
print("Is child (<13)?", is_child)
print("Is exactly 25?", is_exactly_25)
```

Output:
```
Age: 25
Is adult (≥18)? True
Is child (<13)? False
Is exactly 25? True
```

```python
# Comparison operators
a = 10
b = 20

print("")
print("a == b (equal):", a == b)
print("a != b (not equal):", a != b)
print("a < b (less than):", a < b)
print("a > b (greater than):", a > b)
print("a <= b (less or equal):", a <= b)
print("a >= b (greater or equal):", a >= b)
```

Output:
```
a == b (equal): False
a != b (not equal): True
a < b (less than): True
a > b (greater than): False
a <= b (less or equal): True
a >= b (greater or equal): False
```

```python
# Booleans in machine learning
is_trained = False        # Has the model finished training?
use_gpu = True            # Should we use GPU for faster computation?
model_converged = False   # Did training reach optimal point?

print("Model trained?", is_trained)
print("Use GPU?", use_gpu)
print("Converged?", model_converged)

print("\nBooleans help us make decisions in our code!")
```

Output:
```
Model trained? False
Use GPU? True
Converged? False

Booleans help us make decisions in our code!
```

**Key Takeaway**:

- Booleans = `True` or `False` only
- Come from comparisons (`<`, `>`, `==`, etc.)
- Used for decisions and conditions

<div dir="rtl">

### الخلاصة المهمة

</div>

<div dir="rtl">

- Boolean = `True` أو `False` فقط
- تنتج من المقارنات (`<`, `>`, `==`, إلخ)
- تُستخدم لاتخاذ القرارات والشروط في الـ code

</div>


---

### Understanding `print()`

`print()` is how we display information on the screen. Let's explore different ways to use it:

<div dir="rtl">

### فهم الـ `print()`

</div>

<div dir="rtl">

الـ `print()` هي الطريقة التي نعرض بها المعلومات على الشاشة .تعالى نشوف استخدامات ليها:

</div>


```python
# Basic print
print("Hello, World!")
print(42)
print(3.14)
```

Output:
```
Hello, World!
42
3.14
```

```python
# Printing multiple things
name = "Ali"
age = 30

print("Name:", name, "Age:", age)
```

Output:
```
Name: Ali Age: 30
```

```python
# F-strings (formatted strings) - THE BEST WAY!
# Put 'f' before the quotes, then use {variable_name} inside

name = "Mohamed"
age = 25
height = 1.65

print(f"My name is {name}")
print(f"I am {age} years old")
print(f"My height is {height} meters")
print(f"{name} is {age} years old and {height}m tall")
```

Output:
```
My name is Mohamed
I am 25 years old
My height is 1.65 meters
Mohamed is 25 years old and 1.65m tall
```

```python
# F-strings can do math inside!
price = 50
quantity = 3

print(f"Price: ${price}")
print(f"Quantity: {quantity}")
print(f"Total: ${price * quantity}")
```

Output:
```
Price: $50
Quantity: 3
Total: $150
```

```python
# F-strings with formatting
pi = 3.14159265359

print(f"Pi (full): {pi}")
print(f"Pi (2 decimals): {pi:.2f}")  # .2f means 2 decimal places
print(f"Pi (4 decimals): {pi:.4f}")
```

Output:
```
Pi (full): 3.14159265359
Pi (2 decimals): 3.14
Pi (4 decimals): 3.1416
```

---

### Summary: All Data Types Together

Let's see them all in one place:

<div dir="rtl">

### ملخص: كل أنواع الـ Data معاً

</div>

<div dir="rtl">

لنرى كل الأنواع في مكان واحد:

</div>


```python
# A mini ML scenario
model_name = "LinearRegression"      # String
num_epochs = 100                     # Integer
learning_rate = 0.01                 # Float
is_trained = False                   # Boolean
accuracy = 0.95                      # Float

print("╔════════════════════════════════╗")
print("║   Machine Learning Model       ║")
print("╚════════════════════════════════╝")
print(f"Model: {model_name}")
print(f"Epochs: {num_epochs}")
print(f"Learning Rate: {learning_rate}")
print(f"Is Trained: {is_trained}")
print(f"Accuracy: {accuracy * 100}%")
```

Output:
```
╔════════════════════════════════╗
║   Machine Learning Model       ║
╚════════════════════════════════╝
Model: LinearRegression
Epochs: 100
Learning Rate: 0.01
Is Trained: False
Accuracy: 95.0%
```

```python
# Let's check all their types
print("Variable Types:")
print(f"  model_name: {type(model_name).__name__}")  # type is used to show the type of the variable
print(f"  num_epochs: {type(num_epochs).__name__}")
print(f"  learning_rate: {type(learning_rate).__name__}")
print(f"  is_trained: {type(is_trained).__name__}")
print(f"  accuracy: {type(accuracy).__name__}")
```

Output:
```
Variable Types:
  model_name: str
  num_epochs: int
  learning_rate: float
  is_trained: bool
  accuracy: float
```

---

### Type Conversion: Changing Between Types

Sometimes we need to convert data from one type to another. Let's see how:

<div dir="rtl">

### تحويل الأنواع: التغيير بين الأنواع

</div>

<div dir="rtl">

أحياناً نحتاج تحويل البيانات من نوع لآخر. تعالى نشوف
:

</div>


```python
# String to Integer .. casting
age_text = "25"
print(f"Original: '{age_text}' (type: {type(age_text).__name__})")

age_number = int(age_text)
print(f"Converted: {age_number} (type: {type(age_number).__name__})")

# Now we can do math!
print(f"Next year: {age_number + 1}")
```

Output:
```
Original: '25' (type: str)
Converted: 25 (type: int)
Next year: 26
```

```python
# String to Float
price_text = "19.99"
print(f"Original: '{price_text}' (type: {type(price_text).__name__})")

price_number = float(price_text)
print(f"Converted: {price_number} (type: {type(price_number).__name__})")

print(f"With tax: ${price_number * 1.1:.2f}")
```

Output:
```
Original: '19.99' (type: str)
Converted: 19.99 (type: float)
With tax: $21.99
```

```python
# Number to String
count = 42
print(f"Original: {count} (type: {type(count).__name__})")

count_text = str(count)
print(f"Converted: '{count_text}' (type: {type(count_text).__name__})")

# Now we can combine it with other text
message = "You have " + count_text + " items"
print(message)
```

Output:
```
Original: 42 (type: int)
Converted: '42' (type: str)
You have 42 items
```

```python
# Float to Integer (loses decimal part!)
pi = 3.14159
print(f"Original: {pi}")

pi_as_int = int(pi)
print(f"As integer: {pi_as_int}")
print("⚠️ Warning: Decimal part was lost!")
```

Output:
```
Original: 3.14159
As integer: 3
⚠️ Warning: Decimal part was lost!
```

---

### Now It's Your Turn! 🎯

Let's practice what you learned:

<div dir="rtl">

### دورك! 🎯


### Exercise 1.1: Create Variables


<!-- grade_id: ex1_1, points: 0, grade: false, solution: true -->
```python
# Create these variables:
# - weight: a float equal to 0.5
# - bias: a float equal to 0.0
# - num_features: an integer equal to 10
# - dataset_name: a string equal to "mnist"

# YOUR CODE HERE
weight = 0.5 # indentation
bias = 0.0
num_features = 10
dataset_name = "mnist"


print(f"Weight: {weight}")
print(f"Bias: {bias}")
print(f"Number of Features: {num_features}")
print(f"Dataset: {dataset_name}")
```

Output:
```
Weight: 0.5
Bias: 0.0
Number of Features: 10
Dataset: mnist
```

<!-- grade_id: ex1_1_test, points: 1, grade: true, solution: false -->
```python
assert weight == 0.5 and isinstance(weight, float)
assert bias == 0.0 and isinstance(bias, float)
assert num_features == 10 and isinstance(num_features, int)
assert dataset_name == "mnist" and isinstance(dataset_name, str)
print("✓ Exercise 1.1 passed!")
```

Output:
```
✓ Exercise 1.1 passed!
```

### Exercise 1.2: Type Conversion


```python
print(isinstance(weight , float))
```

Output:
```
True
```

<!-- grade_id: ex1_2, points: 0, grade: false, solution: true -->
```python
# Given these strings:
epochs_str = "500"
lr_str = "0.001"

# Convert them to numbers:
# - epochs_int: convert epochs_str to integer
# - lr_float: convert lr_str to float

# YOUR CODE HERE
epochs_int= int(epochs_str)
lr_float = float(lr_str)

print(f"Epochs: {epochs_int} (type: {type(epochs_int).__name__})")
print(f"Learning Rate: {lr_float} (type: {type(lr_float).__name__})")
```

Output:
```
Epochs: 500 (type: int)
Learning Rate: 0.001 (type: float)
```

<!-- grade_id: ex1_2_test, points: 1, grade: true, solution: false -->
```python
assert epochs_int == 500 and isinstance(epochs_int, int)
assert lr_float == 0.001 and isinstance(lr_float, float)
print("✓ Exercise 1.2 passed!")
```

Output:
```
✓ Exercise 1.2 passed!
```

---

## Part 2: Lists - Organizing Multiple Values

### What is a List?

A **list** is a container that holds multiple values in order. Think of it like:

- A shopping list
- A row in a spreadsheet

<div dir="rtl">

### ما هي الـ List؟

</div>

<div dir="rtl">

الـ **list** هي وعاء يحتوي على قيم متعددة بترتيب معين. تخيلها مثل:

</div>

<div dir="rtl">

- قائمة تسوق
- صف في جدول بيانات

</div>

**Visual representation**:

```
fruits = ["apple", "banana", "cherry"]

┌─────────┬─────────┬─────────┐
│ "apple" │"banana" │"cherry" │
└─────────┴─────────┴─────────┘
 Index 0   Index 1   Index 2
```

**Important**: In python, we start counting from 0, not 1!

<div dir="rtl">

**مهم**: نبدأ العد من 0، مش من 1!

</div>

Let's create and explore lists:


```python
# Creating lists
fruits = ["apple", "banana", "cherry"] # square brackets
numbers = [10, 20, 30, 40, 50]
mixed = ["hello", 42, 3.14, True]  # Can mix types!

print("Fruits:", fruits)
print("Numbers:", numbers)
print("Mixed:", mixed)
```

Output:
```
Fruits: ['apple', 'banana', 'cherry']
Numbers: [10, 20, 30, 40, 50]
Mixed: ['hello', 42, 3.14, True]
```

```python
# Empty list
empty_list = ["hello", 42, 3.14, True]
print("Empty list:", empty_list)
print("Length:", len(empty_list)) # length
```

Output:
```
Empty list: ['hello', 42, 3.14, True]
Length: 4
```

### Accessing List Elements (Indexing)

Use square brackets `[]` with the position number (index):

<div dir="rtl">

### الوصول لعناصر الـ List (Indexing)

</div>

<div dir="rtl">

استخدم الأقواس المربعة `[]` مع رقم الموقع (index):

</div>


```python
fruits = ["apple", "banana", "cherry", "date", "elderberry"]

# ()	Parentheses
# []	Square brackets
# {}	Curly braces
# <>	Angle brackets

print("Full list:", fruits)
print("")
print("fruits[0]:", fruits[0])  # First item
print("fruits[1]:", fruits[1])  # Second item
print("fruits[2]:", fruits[2])  # Third item
print("fruits[4]:", fruits[4])  # Fifth item
```

Output:
```
Full list: ['apple', 'banana', 'cherry', 'date', 'elderberry']

fruits[0]: apple
fruits[1]: banana
fruits[2]: cherry
fruits[4]: elderberry
```

```python
# Negative indexing - count from the end!
print("Negative indexing:")
print("fruits[-1]:", fruits[-1])  # Last item
print("fruits[-2]:", fruits[-2])  # Second-to-last
print("fruits[-3]:", fruits[-3])  # Third-to-last
```

Output:
```
Negative indexing:
fruits[-1]: elderberry
fruits[-2]: date
fruits[-3]: cherry
```

```python
# Visual diagram
print("Visual representation:")
print("")
print("Indices:      0         1         2        3           4")
print("           ┌─────────┬─────────┬─────────┬─────────┬────────────┐")
print("Values:    │ 'apple' │'banana' │'cherry' │ 'date'  │'elderberry'│")
print("           └─────────┴─────────┴─────────┴─────────┴────────────┘")
print("Negative:     -5        -4        -3       -2          -1")
```

Output:
```
Visual representation:

Indices:      0         1         2        3           4
           ┌─────────┬─────────┬─────────┬─────────┬────────────┐
Values:    │ 'apple' │'banana' │'cherry' │ 'date'  │'elderberry'│
           └─────────┴─────────┴─────────┴─────────┴────────────┘
Negative:     -5        -4        -3       -2          -1
```

```python
# Real ML example: a data point with multiple features
person = [175, 70, 25]  # [height_cm, weight_kg, age_years]

print("Person data:", person)
print(f"Height: {person[0]} cm")
print(f"Weight: {person[1]} kg")
print(f"Age: {person[2]} years")
```

Output:
```
Person data: [175, 70, 25]
Height: 175 cm
Weight: 70 kg
Age: 25 years
```

### List Slicing - Getting Multiple Items

**Syntax**: `list[start:end]`

- Includes `start`
- **Excludes** `end` (stops before it)

<div dir="rtl">

### تقطيع الـ List - الحصول على عدة عناصر

</div>

<div dir="rtl">

**الصيغة**: `list[start:end]`

</div>

<div dir="rtl">

- يشمل `start`
- **لا يشمل** `end` (يتوقف قبله)

</div>

Let's see it in action:


```python
numbers = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]

print("Full list:", numbers)
print("")
print("numbers[0:3]:", numbers[0:3])    # Items 0, 1, 2 (stops before 3!)
print("numbers[2:5]:", numbers[2:5])    # Items 2, 3, 4
print("numbers[5:8]:", numbers[5:8])    # Items 5, 6, 7
```

Output:
```
Full list: [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]

numbers[0:3]: [0, 1, 2]
numbers[2:5]: [2, 3, 4]
numbers[5:8]: [5, 6, 7]
```

```python
# Shortcuts
print("Shortcuts:")
print("numbers[:3]:", numbers[:3])      # From start to index 3
print("numbers[7:]:", numbers[7:])      # From index 7 to end
print("numbers[-3:]:", numbers[-3:])    # Last 3 items
print("numbers[:-2]:", numbers[:-2])    # Everything except last 2
```

Output:
```
Shortcuts:
numbers[:3]: [0, 1, 2]
numbers[7:]: [7, 8, 9]
numbers[-3:]: [7, 8, 9]
numbers[:-2]: [0, 1, 2, 3, 4, 5, 6, 7]
```

```python
# Visual demonstration
print("Visual example of slicing:")
print("")
print("Index:   0    1    2    3    4    5    6    7    8    9")
print("Value: [ 0,   1,   2,   3,   4,   5,   6,   7,   8,   9 ]")
print("")
print("numbers[2:7] means:")
print("        start here↓              stop before↓")
print("Value: [ 0,   1,   2,   3,   4,   5,   6,   7,   8,   9 ]")
print("                  └──────────────────────┘")
print("Result:", numbers[2:7])
```

Output:
```
Visual example of slicing:

Index:   0    1    2    3    4    5    6    7    8    9
Value: [ 0,   1,   2,   3,   4,   5,   6,   7,   8,   9 ]

numbers[2:7] means:
        start here↓              stop before↓
Value: [ 0,   1,   2,   3,   4,   5,   6,   7,   8,   9 ]
                  └──────────────────────┘
Result: [2, 3, 4, 5, 6]
```

### Modifying Lists

Lists are **mutable** (changeable). Let's see how:

<div dir="rtl">

### تعديل الـ Lists

</div>

<div dir="rtl">

الـ Lists **قابلة للتعديل** (mutable). تعالى نشوف:

</div>


```python
# Changing an element
fruits = ["apple", "banana", "cherry"]
print("Original:", fruits)

fruits[1] = "blueberry"
print("After change:", fruits)
```

Output:
```
Original: ['apple', 'banana', 'cherry']
After change: ['apple', 'blueberry', 'cherry']
```

```python
# Adding elements
numbers = [1, 2, 3]
print("Start:", numbers)

numbers.append(4)      # Add to end
print("After append(4):", numbers)

numbers.append(5)
print("After append(5):", numbers)
```

Output:
```
Start: [1, 2, 3]
After append(4): [1, 2, 3, 4]
After append(5): [1, 2, 3, 4, 5]
```

```python
# Useful list operations
values = [10, 5, 20, 15, 30]

print("List:", values)
print("Length:", len(values))
print("Sum:", sum(values))
print("Min:", min(values))
print("Max:", max(values))
print("Average:", sum(values) / len(values))
```

Output:
```
List: [10, 5, 20, 15, 30]
Length: 5
Sum: 80
Min: 5
Max: 30
Average: 16.0
```

### Nested Lists (2D Data)

Lists can contain other lists! This is like a table or spreadsheet.

<div dir="rtl">

### Lists متداخلة (بيانات ثنائية الأبعاد)

</div>

<div dir="rtl">

الـ Lists ممكن تحتوي على lists أخرى! زي مثلا
جدول أو spreadsheet.

</div>

```
data = [
    [170, 68],  ← Row 0: Person 1 (height, weight)
    [165, 62],  ← Row 1: Person 2
    [180, 75]   ← Row 2: Person 3
]

Like a spreadsheet:
╔════════╦════════╦════════╗
║ Person ║ Height ║ Weight ║
╠════════╬════════╬════════╣
║   0    ║  170   ║   68   ║
║   1    ║  165   ║   62   ║
║   2    ║  180   ║   75   ║
╚════════╩════════╩════════╝
```


```python
# Creating a 2D list
data = [
    [170, 68],
    [165, 62],
    [180, 75]
]

print("Full dataset:")
print(data)
```

Output:
```
Full dataset:
[[170, 68], [165, 62], [180, 75]]
```

```python
# Accessing rows
print("Row 0 (Person 0):", data[0])
print("Row 1 (Person 1):", data[1])
print("Row 2 (Person 2):", data[2])
```

Output:
```
Row 0 (Person 0): [170, 68]
Row 1 (Person 1): [165, 62]
Row 2 (Person 2): [180, 75]
```

```python
# Accessing individual elements
# Syntax: list[row][column]

print("Person 0's height (row 0, col 0):", data[0][0])
print("Person 0's weight (row 0, col 1):", data[0][1])
print("Person 1's height (row 1, col 0):", data[1][0])
print("Person 2's weight (row 2, col 1):", data[2][1])
```

Output:
```
Person 0's height (row 0, col 0): 170
Person 0's weight (row 0, col 1): 68
Person 1's height (row 1, col 0): 165
Person 2's weight (row 2, col 1): 75
```

```python
# How to read data[1][1]:
print("Understanding data[1][1]:")
print("  Step 1: data[1] =>", data[1])
print("  Step 2: [1] from that =>", data[1][1])
print("  Result: Second person's weight")
```

Output:
```
Understanding data[1][1]:
  Step 1: data[1] => [165, 62]
  Step 2: [1] from that => 62
  Result: Second person's weight
```

### Practice Time! 🎯

<div dir="rtl">

### وقت التمرين! 🎯

</div>


### Exercise 2.1: List Basics


<!-- grade_id: ex2_1, points: 0, grade: false, solution: true -->
```python
# Create a list called 'predictions' with these values: 0, 1, 1, 0, 1, 0, 1, 1
# Then extract:
# - first_pred: the first element
# - last_pred: the last element (use negative indexing!)
# - third_pred: the third element

# YOUR CODE HERE
predictions = [0,1,1,0,1,0,1,1]
first_pred =predictions[0]
last_pred = predictions[-1]
third_pred = predictions[2]

print(f"Predictions: {predictions}")
print(f"First: {first_pred}, Third: {third_pred}, Last: {last_pred}")
```

Output:
```
Predictions: [0, 1, 1, 0, 1, 0, 1, 1]
First: 0, Third: 1, Last: 1
```

<!-- grade_id: ex2_1_test, points: 1, grade: true, solution: false -->
```python
assert predictions == [0, 1, 1, 0, 1, 0, 1, 1]
assert first_pred == 0 and third_pred == 1 and last_pred == 1
print("✓ Exercise 2.1 passed!")
```

Output:
```
✓ Exercise 2.1 passed!
```

### Exercise 2.2: List Slicing


<!-- grade_id: ex2_2, points: 0, grade: false, solution: true -->
```python
# Given:
losses = [0.8, 0.65, 0.52, 0.43, 0.38, 0.34, 0.31, 0.29, 0.27, 0.26]

# Extract:
# - first_three: first 3 values
# - last_three: last 3 values
# - middle: values from index 3 to 6 (inclusive of 3, exclusive of 7)

# YOUR CODE HERE
first_three = losses[:3]
last_three = losses[-3:]
middle = losses[3:7]

print(f"First three: {first_three}")
print(f"Last three: {last_three}")
print(f"Middle: {middle}")
```

Output:
```
First three: [0.8, 0.65, 0.52]
Last three: [0.29, 0.27, 0.26]
Middle: [0.43, 0.38, 0.34, 0.31]
```

<!-- grade_id: ex2_2_test, points: 1, grade: true, solution: false -->
```python
assert first_three == [0.8, 0.65, 0.52]
assert last_three == [0.29, 0.27, 0.26]
assert middle == [0.43, 0.38, 0.34, 0.31]
print("✓ Exercise 2.2 passed!")
```

Output:
```
✓ Exercise 2.2 passed!
```

### Exercise 2.3: List Operations


<!-- grade_id: ex2_3, points: 0, grade: false, solution: true -->
```python
# Start with:
feature_values = [1.2, 3.4, 2.1]

# Do these operations:
# 1. Append 4.5 to the list
# 2. Calculate length → store in 'num_features'
# 3. Calculate sum → store in 'total'

# YOUR CODE HERE
feature_values.append(4.5)
num_features = len(feature_values)
total = sum(feature_values)

print(f"Features: {feature_values}")
print(f"Count: {num_features}")
print(f"Total: {total}")
```

Output:
```
Features: [1.2, 3.4, 2.1, 4.5]
Count: 4
Total: 11.2
```

<!-- grade_id: ex2_3_test, points: 1, grade: true, solution: false -->
```python
assert feature_values == [1.2, 3.4, 2.1, 4.5]
assert num_features == 4 and abs(total - 11.2) < 0.01
print("✓ Exercise 2.3 passed!")
```

Output:
```
✓ Exercise 2.3 passed!
```

### Exercise 2.4: 2D Lists


<!-- grade_id: ex2_4, points: 0, grade: false, solution: true -->
```python
# Create a 2D list 'data' with 3 rows:
# Row 0: [170, 68]
# Row 1: [165, 62]
# Row 2: [180, 75]
#
# Then extract:
# - first_sample: entire first row
# - second_weight: weight from second row (index 1, then 1)

# YOUR CODE HERE
data =[ [170, 68],
        [165 , 62],
         [180, 75]]
first_sample = data[0]
second_weight = data[1][1]

print(f"Data: {data}")
print(f"First sample: {first_sample}")
print(f"Second weight: {second_weight}")
```

Output:
```
Data: [[170, 68], [165, 62], [180, 75]]
First sample: [170, 68]
Second weight: 62
```

<!-- grade_id: ex2_4_test, points: 1, grade: true, solution: false -->
```python
assert data == [[170, 68], [165, 62], [180, 75]]
assert first_sample == [170, 68] and second_weight == 62
print("✓ Exercise 2.4 passed!")
```

Output:
```
✓ Exercise 2.4 passed!
```

---

## Part 3: Loops - Repeating Actions

### What is a Loop?

A **loop** repeats the same action multiple times automatically.

**Real-world examples**:

- "For each email in my inbox, check if it's spam"
- "For each student in class, grade their homework"

**Why loops matter**: Instead of writing the same code 100 times, write it once and loop 100 times!

<div dir="rtl">

### ما هي الـ Loop؟

</div>

<div dir="rtl">

الـ **loop** تكرر نفس العملية عدة مرات تلقائياً.

</div>

<div dir="rtl">

** أمثلة **:

</div>

<div dir="rtl">

- "لكل email  شوف إذا كان spam"
- "لكل طالب في الفصل، صحح واجبه"

</div>

<div dir="rtl">

**لماذا الـ loops مهمة**: بدلاً من كتابة نفس الـ code 100 مرة، اكتبه مرة واحدة وكرره 100 مرة!

</div>

### The `for` Loop


```python
# Without a loop (repetitive!)
print("Without loop:")
print("Hello 1")
print("Hello 2")
print("Hello 3")
```

```python
# With a loop (elegant!)
print("With loop:")
for i in range(3):
  print(f"Hello {i}")
```

Output:
```
With loop:
Hello 0
Hello 1
Hello 2
```

```python
# Understanding range()
print("range(5) produces:")
for number in range(5):
    print(number)  # Starts at 0, ends at 4
```

Output:
```
range(5) produces:
0
1
2
3
4
```

```python
# Looping through a list
fruits = ["apple", "banana", "cherry"]

print("Method 1: Loop through items directly")
for fruit in fruits:
    print(f"I like {fruit}")
```

Output:
```
Method 1: Loop through items directly
I like apple
I like banana
I like cherry
```

```python
# Method 2: Loop through indices
print("Method 2: Loop through indices")
for i in range(len(fruits)):
    print(f"Index {i}: {fruits[i]}")
```

Output:
```
Method 2: Loop through indices
Index 0: apple
Index 1: banana
Index 2: cherry
```

```python
# Accumulating values in a loop
numbers = [10, 20, 30, 40, 50]
total = 0

print("Calculating sum step by step:")
for num in numbers:
    total = total + num
    print(f"  Added {num}, running total: {total}")

print(f"\nFinal total: {total}")
```

Output:
```
Calculating sum step by step:
  Added 10, running total: 10
  Added 20, running total: 30
  Added 30, running total: 60
  Added 40, running total: 100
  Added 50, running total: 150

Final total: 150
```

```python
# ML Example: Training simulation
print("=== Training a Model ===")
initial_loss = 1.0

for epoch in range(5):
    # Simulate loss decreasing
    current_loss = initial_loss * (0.7 ** epoch)
    print(f"Epoch {epoch + 1}: Loss = {current_loss:.4f}")

print("\nTraining complete!")
```

Output:
```
=== Training a Model ===
Epoch 1: Loss = 1.0000
Epoch 2: Loss = 0.7000
Epoch 3: Loss = 0.4900
Epoch 4: Loss = 0.3430
Epoch 5: Loss = 0.2401

Training complete!
```

### Loops with Conditions

We can use `if` to make decisions:

<div dir="rtl">

### Loops مع شروط

</div>

<div dir="rtl">

يمكننا استخدام `if` لاتخاذ قرارات:

</div>


```python
# Check each number
numbers = [1, 5, 8, 12, 15, 20]

print("Checking which numbers are > 10:")
for num in numbers:
    if num > 10:
        print(f"  {num} is greater than 10 ✓")
    else:
        print(f"  {num} is not greater than 10")
```

```python
# Counting specific items
predictions = [0, 1, 1, 0, 1, 0, 1, 1]
count_ones = 0

for pred in predictions:
    if pred == 1:
        count_ones += 1  # Same as: count_ones = count_ones + 1

print(f"Predictions: {predictions}")
print(f"Number of 1s: {count_ones}")
print(f"Number of 0s: {len(predictions) - count_ones}")
```

Output:
```
Predictions: [0, 1, 1, 0, 1, 0, 1, 1]
Number of 1s: 5
Number of 0s: 3
```

### Nested Loops

Loops inside loops - useful for 2D data or even higher dimensional data:

<div dir="rtl">

### Loops متداخلة

</div>

<div dir="rtl">

Loops داخل loops - مفيدة للبيانات ثنائية الأبعاد أو أكثر
:

</div>


```python
# Simple nested loop
print("Multiplication table (3x3):")
for i in range(1, 4):
    for j in range(1, 4):
        print(f"{i} × {j} = {i*j}")
    print("---")  # Separator after each row

```

Output:
```
Multiplication table (3x3):
1 × 1 = 1
1 × 2 = 2
1 × 3 = 3
---
2 × 1 = 2
2 × 2 = 4
2 × 3 = 6
---
3 × 1 = 3
3 × 2 = 6
3 × 3 = 9
---
```

```python
# Processing 2D data
data = [
    [1, 2, 3],
    [4, 5, 6],
    [7, 8, 9]
]

print("Processing each element:")
for row_idx, row in enumerate(data):
    for col_idx, value in enumerate(row):
        print(f"  Position [{row_idx}][{col_idx}] = {value}")
```

Output:
```
Processing each element:
  Position [0][0] = 1
  Position [0][1] = 2
  Position [0][2] = 3
  Position [1][0] = 4
  Position [1][1] = 5
  Position [1][2] = 6
  Position [2][0] = 7
  Position [2][1] = 8
  Position [2][2] = 9
```

### Practice Time! 🎯

<div dir="rtl">

### وقت التمرين! 🎯

</div>


### Exercise 3.1: Basic Loop


<!-- grade_id: ex3_1, points: 0, grade: false, solution: true -->
```python
predictions = [0, 1, 1, 0, 1]

# Use a loop to create strings like "Sample 0: Class 0"
# Store all strings in output_lines list

# YOUR CODE HERE
output_lines =[]
for idx in range( len(predictions)):
  if (predictions[idx]==0):
    output_lines.append( f"Sample {idx}: Class {0}")
  else :
    output_lines.append(f"Sample {idx}: Class {1}")

print(output_lines)
```

Output:
```
['Sample 0: Class 0', 'Sample 1: Class 1', 'Sample 2: Class 1', 'Sample 3: Class 0', 'Sample 4: Class 1']
```

<!-- grade_id: ex3_1_test, points: 1, grade: true, solution: false -->
```python
expected = ["Sample 0: Class 0", "Sample 1: Class 1", "Sample 2: Class 1", "Sample 3: Class 0", "Sample 4: Class 1"]
assert output_lines == expected
print("✓ Exercise 3.1 passed!")
```

Output:
```
✓ Exercise 3.1 passed!
```

### Exercise 3.2: Counting with Loops


<!-- grade_id: ex3_2, points: 0, grade: false, solution: true -->
```python
predictions = [0, 1, 1, 0, 1, 0, 1, 1]
true_labels = [0, 1, 0, 0, 1, 0, 1, 1]

# Count correct predictions
# Store in correct_count

# YOUR CODE HERE
correct_count = 0
for idx in range (len(predictions)):
  if (predictions[idx] == true_labels[idx]):
    correct_count +=1
print(f"Correct: {correct_count}/{len(predictions)}")
print(f"Accuracy: {correct_count/len(predictions)*100:.1f}%")
```

Output:
```
Correct: 7/8
Accuracy: 87.5%
```

<!-- grade_id: ex3_2_test, points: 1, grade: true, solution: false -->
```python
assert correct_count == 7
print("✓ Exercise 3.2 passed!")
```

Output:
```
✓ Exercise 3.2 passed!
```

### Exercise 3.3: Calculate Column Means


<!-- grade_id: ex3_3, points: 0, grade: false, solution: true -->
```python
# Data: 3 people with [height, weight]
data = [
    [170, 68],
    [165, 62],
    [180, 75]
]

# Calculate mean of each column
# mean_height: average of first column
# mean_weight: average of second column

# YOUR CODE HERE
col_weight = 1
col_hight = 0
sum_hight = 0
sum_weight = 0
for row in range( len(data)):
  sum_hight += data[row][col_hight]
  sum_weight += data[row][col_weight]
mean_height = sum_hight/len(data)
mean_weight = sum_weight/len(data)


print(f"Mean height: {mean_height:.2f} cm")
print(f"Mean weight: {mean_weight:.2f} kg")
```

Output:
```
Mean height: 171.67 cm
Mean weight: 68.33 kg
```

<!-- grade_id: ex3_3_test, points: 1, grade: true, solution: false -->
```python
assert abs(mean_height - 171.67) < 0.1 and abs(mean_weight - 68.33) < 0.1
print("✓ Exercise 3.3 passed!")
```

Output:
```
✓ Exercise 3.3 passed!
```

### Exercise 3.4: Nested Loop


<!-- grade_id: ex3_4, points: 0, grade: false, solution: true -->
```python
data = [
    [1, 2, 3],
    [4, 5, 6],
    [7, 8, 9]
]

# Sum ALL values using nested loops
# Store in total_sum

# YOUR CODE HERE
total_sum = 0
for row in range (len(data)):
  for col in range (len(data[row])):
    total_sum += data[row][col]

print(f"Total sum: {total_sum}")
```

Output:
```
Total sum: 45
```

<!-- grade_id: ex3_4_test, points: 1, grade: true, solution: false -->
```python
assert total_sum == 45
print("✓ Exercise 3.4 passed!")
```

Output:
```
✓ Exercise 3.4 passed!
```

---

## Part 4: Functions - Reusable Code Blocks

### What is a Function?

A **function** is like a recipe:

- Give it ingredients (inputs)
- It follows steps (code)
- Returns a dish (output)

**Why use functions?**

1. **Reusability**: Write once, use many times
2. **Organization**: Keep code tidy
3. **Easier to test**: Test small pieces

<div dir="rtl">

### ما هي الـ Function؟

</div>

<div dir="rtl">

الـ **function** مثل وصفة طبخ:

</div>

<div dir="rtl">

- تعطيها مكونات (inputs)
- تتبع خطوات (code)
- تعيد لك طبق (output)

</div>

<div dir="rtl">

**لماذا نستخدم functions؟**

</div>

<div dir="rtl">

1. **إعادة الاستخدام**: اكتب مرة، استخدم مرات كثيرة
2. **التنظيم**: إبقاء الـ code منظم
3. **أسهل للاختبار**: اختبر أجزاء صغيرة

</div>

### Creating Functions:


```python
# Simple function
def greet(name): # parameter
    """Print a greeting message."""
    print(f"Hello, {name}!")

# Use it .. call
greet("Wael") # argument ..
greet("Yahia")
greet("Mona")
```

Output:
```
Hello, Wael!
Hello, Yahia!
Hello, Mona!
```

```python
# Function that returns a value
def add(a, b):
    """Add two numbers."""
    result = a + b
    return result


# Use it
answer = add(5, 3)
print(f"5 + 3 = {answer}")

answer2 = add(100, 200)
print(f"100 + 200 = {answer2}")
```

Output:
```
5 + 3 = 8
100 + 200 = 300
```

```python
# Function parts explained
def calculate_average(numbers):
    """Calculate the average of a list."""
    #   ↑ Docstring: explains what function does

    total = sum(numbers)
    count = len(numbers)
    average = total / count

    return average
    #      ↑ Returns the result


# Test it
scores = [85, 90, 78, 92, 88]
avg = calculate_average(scores)
print(f"Scores: {scores}")
print(f"Average: {avg}")
```

Output:
```
Scores: [85, 90, 78, 92, 88]
Average: 86.6
```

```python
# ML Example: Accuracy function
def calculate_accuracy(predictions, true_labels):
    """
    Calculate classification accuracy.

    Args:
        predictions: List of predicted labels
        true_labels: List of true labels

    Returns:
        Accuracy as percentage (0-100)
    """
    correct = 0
    for i in range(len(predictions)):
        if predictions[i] == true_labels[i]:
            correct += 1  # this is the same as correct = correct + 1

    accuracy = (correct / len(predictions)) * 100
    return accuracy


# Test it
pred = [0, 1, 1, 0, 1]
true = [0, 1, 0, 0, 1]
acc = calculate_accuracy(pred, true)
print(f"Accuracy: {acc:.1f}%")
```

Output:
```
Accuracy: 80.0%
```

### More ML Functions:

<div dir="rtl">

### المزيد من functions لل Machine Learning:

</div>


```python
# Euclidean distance
def distance(point1, point2):
    """
    Calculate distance between two 2D points.

    Formula: sqrt((x1-x2)² + (y1-y2)²)
    """
    x1, y1 = point1
    x2, y2 = point2

    dist = ((x1-x2)**2 + (y1-y2)**2)**0.5
    return dist


# Test
p1 = [0, 0]
p2 = [3, 4]
d = distance(p1, p2)
print(f"Distance from {p1} to {p2}: {d}")
```

Output:
```
Distance from [0, 0] to [3, 4]: 5.0
```

```python
# Normalization function (preview this will be discussed in the lecture in detail but see how to implement it here!)
def normalize(values):
    """
    Normalize values to [0, 1] range.

    Formula: (x - min) / (max - min)
    """
    min_val = min(values)
    max_val = max(values)
    range_val = max_val - min_val

    normalized = []
    for val in values:
        norm_val = (val - min_val) / range_val
        normalized.append(norm_val)

    return normalized


# Test
data = [10, 20, 30, 40, 50]
norm_data = normalize(data)
print(f"Original: {data}")
print(f"Normalized: {[f'{x:.2f}' for x in norm_data]}")
```

Output:
```
Original: [10, 20, 30, 40, 50]
Normalized: ['0.00', '0.25', '0.50', '0.75', '1.00']
```

### Practice Time! 🎯

<div dir="rtl">

### وقت التمرين! 🎯

</div>


### Exercise 4.1: Mean Function


<!-- grade_id: ex4_1, points: 0, grade: false, solution: true -->
```python
def calculate_mean(values):
    """
    Calculate mean (average) of a list.

    Args:
        values: List of numbers

    Returns:
        Mean as float
    """
    # YOUR CODE HERE
    sum = 0
    for val in values:
      sum += val
    return sum/len(values)


# Test
test = [10, 20, 30, 40, 50]
result = calculate_mean(test)
print(f"Mean of {test}: {result}")
```

Output:
```
Mean of [10, 20, 30, 40, 50]: 30.0
```

<!-- grade_id: ex4_1_test, points: 1, grade: true, solution: false -->
```python
assert abs(calculate_mean([1, 2, 3, 4, 5]) - 3.0) < 0.01
assert abs(calculate_mean([10, 20, 30]) - 20.0) < 0.01
print("✓ Exercise 4.1 passed!")
```

Output:
```
✓ Exercise 4.1 passed!
```

### Exercise 4.2: Distance Function


<!-- grade_id: ex4_2, points: 0, grade: false, solution: true -->
```python
def euclidean_distance(point1, point2):
    """
    Calculate Euclidean distance between two 2D points.

    Formula: sqrt((x1-x2)² + (y1-y2)²)
    """
    # YOUR CODE HERE
    x1 ,y1 = point1
    x2 ,y2 = point2
    dist = ((x1-x2)**2+(y1-y2)**2)**0.5
    return dist

# Test
p1, p2 = [0, 0], [3, 4]
dist = euclidean_distance(p1, p2)
print(f"Distance: {dist}")
```

Output:
```
Distance: 5.0
```

<!-- grade_id: ex4_2_test, points: 1, grade: true, solution: false -->
```python
assert abs(euclidean_distance([0, 0], [3, 4]) - 5.0) < 0.01
assert abs(euclidean_distance([0, 0], [1, 1]) - 1.414) < 0.01
print("✓ Exercise 4.2 passed!")
```

Output:
```
✓ Exercise 4.2 passed!
```

### Exercise 4.3: Normalization Function


<!-- grade_id: ex4_3, points: 0, grade: false, solution: true -->
```python
def normalize_list(values):
    """
    Normalize to [0,1] using min-max scaling.

    Formula: (x - min) / (max - min)
    """
    # YOUR CODE HERE
    min_val = min(values)
    max_val = max(values)
    range_val = max_val - min_val
    normalized = []
    for val in values:
      norm_val = (val - min_val) / range_val
      normalized.append(norm_val)
    return normalized


# Test
data = [10, 20, 30, 40, 50]
result = normalize_list(data)
print(f"Normalized: {[f'{x:.2f}' for x in result]}")
```

Output:
```
Normalized: ['0.00', '0.25', '0.50', '0.75', '1.00']
```

<!-- grade_id: ex4_3_test, points: 1, grade: true, solution: false -->
```python
r = normalize_list([0, 50, 100])
assert abs(r[0]-0.0) < 0.01 and abs(r[1]-0.5) < 0.01 and abs(r[2]-1.0) < 0.01
print("✓ Exercise 4.3 passed!")
```

Output:
```
✓ Exercise 4.3 passed!
```

### Exercise 4.4: Sigmoid Function


<!-- grade_id: ex4_4, points: 0, grade: false, solution: true -->
```python
def simple_sigmoid(x):
    """
    Sigmoid: σ(x) = 1 / (1 + e^(-x))
    """
    # YOUR CODE HERE
    sigmoid = 1/(1+2.71828**(-x))
    return sigmoid


# Test
for x in [-5, 0, 5]:
    print(f"sigmoid({x}) = {simple_sigmoid(x):.4f}")
```

Output:
```
sigmoid(-5) = 0.0067
sigmoid(0) = 0.5000
sigmoid(5) = 0.9933
```

<!-- grade_id: ex4_4_test, points: 1, grade: true, solution: false -->
```python
assert abs(simple_sigmoid(0) - 0.5) < 0.01
assert simple_sigmoid(10) > 0.9 and simple_sigmoid(-10) < 0.1
print("✓ Exercise 4.4 passed!")
```

Output:
```
✓ Exercise 4.4 passed!
```

---

## Final Challenge: Complete ML Pipeline without a model

Combine everything you learned!

<div dir="rtl">

## التحدي النهائي: بناء ML Pipeline كامل بدون موديل

</div>

<div dir="rtl">

اجمع كل ما تعلمته!

</div>


<!-- grade_id: final, points: 0, grade: false, solution: true -->
```python
# Given data
data_points = [1.5, 4.2, 2.8, 5.6, 3.1]
true_labels = [0, 1, 0, 1, 1]

# Task 1: Make predictions (if > 3.0, predict 1, else 0)
# YOUR CODE HERE
def make_predictions(data_points , true_lables):
  predictions =[]
  for point in data_points :
    if ( point > 3.0 ):
      predictions.append(1)
    else :
      predictions.append(0)
  return predictions
predictions = make_predictions(data_points , true_labels)


# Task 2: Calculate accuracy
# YOUR CODE HERE
def calc_accuracy (predictions , true_lables):
  correct = 0.0
  for idx in range (len(predictions)):
    if (predictions[idx] == true_lables[idx]):
      correct +=1
  return (correct/len(predictions))*100
accuracy = calc_accuracy(predictions , true_labels)

# Task 3: Normalize data
# YOUR CODE HERE
def normalize_list (data_points):
  min_val = min(data_points)
  max_val = max(data_points)
  range = max_val - min_val
  normalized_list = []
  for point in data_points :
    norm_val = (point - min_val)/range
    normalized_list.append(norm_val)
  return normalized_list
normalized_data = normalize_list(data_points)

# Results
print("╔═══════════════════════════════╗")
print("║   ML Pipeline Results        ║")
print("╚═══════════════════════════════╝")
print(f"\nData: {data_points}")
print(f"Normalized: {[f'{x:.3f}' for x in normalized_data]}")
print(f"\nTrue: {true_labels}")
print(f"Pred: {predictions}")
print(f"\nAccuracy: {accuracy:.1f}%")
```

Output:
```
╔═══════════════════════════════╗
║   ML Pipeline Results        ║
╚═══════════════════════════════╝

Data: [1.5, 4.2, 2.8, 5.6, 3.1]
Normalized: ['0.000', '0.659', '0.317', '1.000', '0.390']

True: [0, 1, 0, 1, 1]
Pred: [0, 1, 0, 1, 1]

Accuracy: 100.0%
```

<!-- grade_id: final_test, points: 2, grade: true, solution: false -->
```python
assert predictions == [0, 1, 0, 1, 1] and accuracy == 100.0
assert len(normalized_data) == 5 and all(0 <= x <= 1 for x in normalized_data)
print("✓ Final challenge passed!")
print("\n🎉 You built a complete ML pipeline!")
```

Output:
```
✓ Final challenge passed!

🎉 You built a complete ML pipeline!
```

---

## Congratulations! 🎉

You've completed Week 1!

<div dir="rtl">

## مبروك! 🎉

</div>

<div dir="rtl">

أنت كملت الأسبوع الأول!

</div>

### What You Learned:

✅ **Variables & Types** (int, float, str, bool)  
✅ **Lists** (indexing, slicing, 2D lists)  
✅ **Loops** (for, range, nested loops)  
✅ **Functions** (defining, calling, returning values)

<div dir="rtl">

### ما تعلمته:

</div>

<div dir="rtl">

✅ **Variables و Types** (int, float, str, bool)  
✅ **Lists** (indexing, slicing, lists ثنائية الأبعاد)  
✅ **Loops** (for, range, loops متداخلة)  
✅ **Functions** (تعريف، استدعاء، إرجاع قيم)

</div>

### Next Week:

🔹 NumPy arrays (supercharged lists)  
🔹 Vectorization (fast ML operations)  
🔹 Matplotlib (data visualization)  
🔹 Object-Oriented Programming

<div dir="rtl">

### الأسبوع القادم:

</div>

<div dir="rtl">

🔹 NumPy arrays (lists محسّنة)  
🔹 Vectorization (عمليات ML سريعة)  
🔹 Matplotlib (تصوير البيانات)

</div>
🔹 Object-Oriented Programming

### Week 3 Preview:

You'll implement **Linear Regression** - your first real ML algorithm! 🚀

<div dir="rtl">

### نظرة على الأسبوع 3:

</div>

<div dir="rtl">

ستنفذ **Linear Regression** - أول خوارزمية ML حقيقية! 🚀

</div>

**Submission**: Save notebook, verify all ✓ marks, submit to Dr. Mahmoud for grading.

<div dir="rtl">

**التسليم**: احفظ الـ notebook، تأكد من كل علامات ✓، سلمها إلى د محمود للتصحيح

</div>

**See you next week!** 🎓

<div dir="rtl">

**نراكم الأسبوع القادم!** 🎓

</div>
